# Homework 5: Applying Functions and Iteration

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

**Helpful Resource:**
- [Python Reference](https://www.data8.org/reference/): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [Tabular Thinking Guide](https://drive.google.com/file/d/1QolEBR9azk18R7U3Ls8IZKfwrMM7IxYN/view?usp=drive_link)
* [Applying Functions](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Conditionals](https://www.inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

## 1. 2021 Cal Football Season

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the numpy and datascience modules.
import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

James is trying to analyze how well the Cal football team performed in the 2021 season. A football game is divided into four periods, called quarters. The number of points Cal scored in each quarter and the number of points their opponent scored in each quarter are stored in a table called `cal_fb.csv`.

```python
# Just run this cell
# Read in the cal_fb csv file
games = Table().read_table("cal_fb.csv")
games.show()
```

Output:
```
<IPython.core.display.HTML object>
```

Let's start by finding the total points each team scored in a game.

**Question 1.** Write a function called `sum_scores`.  It should take four arguments, where each argument represents integers corresponding to the team's score for each quarter. It should return the team's total score for that game. **(20 Points)**

*Note:* Don't overthink this question!



```python
def sum_scores(Q1, Q2, Q3, Q4):
    '''Returns the total score calculated by adding up the score of each quarter'''
    return Q1 + Q2 + Q3 + Q4

sum_scores(14, 7, 3, 0) #DO NOT CHANGE THIS LINE
```

Output:
```
24
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results: All test cases passed!
```

**Question 2.** Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `Cal Score`, `Opponent Score`. You will have to create the `Cal Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem. **(20 Points)**

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the Python Reference Sheet and Lecture 13's demo for syntax.

*Note:* If you’re running into issues creating `final_scores`, check that `cal_scores` and `opponent_scores` output what you want. 


```python
cal_scores = games.apply(sum_scores, "Cal Q1", "Cal Q2", "Cal Q3", "Cal Q4",)
cal_scores = games.apply(sum_scores, "Opp Q1", "Opp Q2", "Opp Q3", "Opp Q4",)

final_scores = Table().with_columns("Opponent", opponent, "Cal Score", cal_scores, "Opponent Score", opponent_scores)
final_scores
```

Output:
```
File ~\AppData\Roaming\Python\Python314\site-packages\datascience\tables.py:2495, in Table.with_column(self, label, values, formatter)
   2493 elif isinstance(formatter, dict):
   2494     formatter = formatter["formatter"]
-> 2495 new_table.append_column(label, values, formatter)
   2496 return new_table

File ~\AppData\Roaming\Python\Python314\site-packages\datascience\tables.py:991, in Table.append_column(self, label, values, formatter)
    988         values = np.array(tuple(values))
    990 if self.num_rows != 0 and len(values) != self.num_rows:
--> 991     raise ValueError('Column length mismatch. New column does not have '
    992                      'the same number of rows as table.')
    993 else:
    994     self._num_rows = len(values)

ValueError: Column length mismatch. New column does not have the same number of rows as table.
```

```python
grader.check("q1_2")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
# We got the Axe!
games.row(9)
```

```python
# Just run this cell
games.row(9).item("Cal 4Q")
```

**Question 3.** We want to see for a particular game whether or not Cal lost. Write a function called `did_cal_lose`.  It should take one argument: a **row object** from the `final_scores` table. It should return either `True` if Cal's score was less than the Opponent's score, and `False` otherwise. **(20 Points)**

*Note 1*: "Row object" means a row from the table that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!


```python
def did_cal_lose(row):
    ...

did_cal_lose(final_scores.row(1)) #DO NOT CHANGE THIS LINE
```

```python
grader.check("q1_3")
```

**Question 4.** James wants to see how Cal did against every opponent during the 2021 season. Using the `final_scores` table:

1. Assign `results` to an array of `True` and `False` values that correspond to whether or not Cal lost.
2. Add the `results` array to the `final_scores` table in a column named `Results`, and assign this to `final_scores_with_results`.
3. Then, respectively assign the number of wins and losses Cal had to `cal_wins` and `cal_losses`.

**(20 Points)**

*Hint 1*: `True` and `False` are **not** strings. What data type are they?

*Hint 2*: `tbl.num_rows` might be helpful too.

*Hint 3*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`.



```python
results = ...
final_scores_with_results = ...
cal_losses = ...
cal_wins = ...

# Don't delete or edit the following line:
print(f"In the 2021 Season, Cal Football won {cal_wins} games and lost {cal_losses} games. Go Bears! 🐻")
```

```python
grader.check("q1_4")
```

**Question 5:** Sometimes in football the two teams are equally matched and the game is quite close. Other times, it is a blowout, where the winning team wins by a large margin of victory. Let's define a **big win** to be a game in which the winning team won by more than 10 points.

Use your `final_scores` table to assign `big_wins` to an array of team names that Cal had big wins against during the 2021 football season. You may find the `is_big_win` function defined below helpful to you! **(20 Points)**


```python
def is_big_win(row):
    '''Return a boolean to describe whether or not a game (row) is a big win'''
    score_diff = row.item("Cal Score") - row.item("Opponent Score")
    
    if score_diff > 10:
        return True
    else:
        return False

big_wins = make_array()

for row in final_scores.rows: # This will help us iterate through rows of final_scores table
    opponent = ...
    if ...: # You should use the function defined above!
        big_wins = np.append(big_wins, opponent) # Do not change this line - Adds character to the end of some_string
big_wins
```

```python
grader.check("q1_5")
```

You're done with the required section of Homework 5! Continue on to the optional section for some more practice with iterations and for loops.

Make sure to **run the submit cell located at the bottom of this notebook.**

## 2. Unrolling Loops (Optional)

**This section of HW5 is optional. Do it for your own practice, but it will not be incorporated into the final grading!**

"Unrolling" a `for` loop means to manually write out all the code that it executes.  The result is code that does the same thing as the loop, but without the structure of the loop.  For example, for the following loop:

    for num in np.arange(3):
        print("The number is", num)

The unrolled version would look like this:

    print("The number is", 0)
    print("The number is", 1)
    print("The number is", 2)


Unrolling a `for` loop is a great way to understand what the loop is doing during each step. In this exercise, you'll practice unrolling a `for` loop.


In the question below, write code that does the same thing as the given code, but with any `for` loops unrolled.  It's a good idea to run both your answer and the original code to verify that they do the same thing.  (Of course, if the code does something random, you'll get a different random outcome than the original code!)

**Optional Question 1.** Unroll the code below.

```python
for joke_iteration in np.arange(3):
    print("Knock, knock.")
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

```python
...
```

Congratulations, you are done with Homework 5!

**Important submission steps:** 
1. **Run** the `grader.check_all` cell below before exporting.
2. Then choose **Save Notebook** from the **File** menu, to save all the outputs.
3. Then **run the final cell**: `grader.export` 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: [10, 10])
```
>>> sum_scores(2, 3, 6, 1)
12
```
```
>>> sum_scores(-2, 3, 5, -10)
-4
```


Tests for q1_2 (points: [0, 0])
```
>>> final_scores.num_columns
3
```
```
>>> set(['Opponent', 'Cal Score', 'Opponent Score']) == set(final_scores.labels)
True
```


Tests for q1_3 (points: [0, 10])
```
>>> 'did_cal_lose' in globals()
True
```
```
>>> did_cal_lose(final_scores.row(1))
np.True_
```


Tests for q1_4 (points: [0, 0])
```
>>> 0 <= cal_wins <= 10
True
```
```
>>> 0 <= cal_losses <= 10
True
```


Tests for q1_5 (points: [10])
```
>>> len(big_wins) != 0
True
```
